# Exercice Sheet: B Expressions

You are given a B machine with three given sets:
* a set ```S``` of online shops
* a set ```A``` of articles that can be ordered
* a set ```C``` of clients
You are also given two relations:
* a relation ```delivers:S<->A``` describing which shops can deliver which articles
* a relation ```orders:C<->POW(A)``` describing current outstanding orders, an order consisting of a set of articles.

```prob
::load
MACHINE exercise1
DEFINITIONS "LibraryHash.def"
SETS
  S = {s1,s2,s3,s4,s5} /* Shops */ ;
  A = {a1,a2,a3,a4,a5} /* Articles that can be ordered */;
  C = {c1,c2,c3} /* Clients */
CONSTANTS delivers, orders
PROPERTIES
  delivers = {s1↦a1,s1↦a2, s2↦a4, s3↦a3, s3↦a2, s4↦a4, s4↦a5} &
  orders = {c1 ↦ {a1,a5}, c1 ↦ {a4}, c2 ↦ {a1,a3,a5}, c3 ↦ {a2,a3}, c3 ↦ {a1}}
END
```

Output:
```
Loaded machine: exercise1
```

```prob
:constants
```

Output:
```
Executed operation: SETUP_CONSTANTS()
```

The above machine also contains particular instantiations of the above sets and relations, which will be used in the exercises below.

## Exercise 1
Write a B expression to compute the set of clients that have ordered article a5. Use the ```:let sol1 EXPRESSION``` in the cell below.

<!-- grade_id: cell-bab176fbf0a75b6f, points: 0, grade: false, solution: true -->
```prob
:let sol1 {c| a5 : union(orders[{c}])}
```

Output:
```
{c1,c2}
```

<!-- grade_id: cell-2b74a8ed88e60492, points: 0, grade: true, solution: false -->
```prob
:assert SHA_HASH_HEX(sol1) = "26365f212fb4cbe0ef2adc9eb42d3c6fac15442b"
```

Output:
```
TRUE
```

## Exercise 2
Write a B expression which computes the set of all articles ordered by at least one client. Use the ```:let sol2 EXPRESSION``` in the cell below.

<!-- grade_id: cell-5fb80f74fa97df15, points: 0, grade: false, solution: true -->
```prob
:let sol2 union(ran(orders))
```

Output:
```
{a1,a2,a3,a4,a5}
```

<!-- grade_id: cell-b1cc285b1a32f2bd, points: 0, grade: true, solution: false -->
```prob
SHA_HASH_HEX(sol2)
```

Output:
```
"29bac35e283e6637fd22c2c8ee4a17a8d4ff52a3"
```

<!-- grade_id: cell-94e08eff5d6d66bf, points: 0, grade: true, solution: false -->
```prob
:assert SHA_HASH_HEX(sol2) = "29bac35e283e6637fd22c2c8ee4a17a8d4ff52a3"
```

Output:
```
TRUE
```

# Comments:
To lock cells you can do this:


* menu: view → cell toolbar → edit metadata.
* make this like
```
{
  "trusted": true,
  "editable": false,
  "deletable": false
}
```